# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    # raise NotImplementedError()
    import pandas as pd
    df = pd.read_csv('NISPUF17.csv', index_col=0)
    # selecting the needed column from the dataframe
    df_mother = df['EDUC1']
    total = df_mother.count()
    less_than_high_school = df_mother[df_mother == 1].count()
    high_school = df_mother[df_mother == 2].count()
    more_than_high_school_not_college = df_mother[df_mother == 3].count()
    college = df_mother[df_mother == 4].count()
    dictionary = {"less than high school": less_than_high_school/total,
                  "high school": high_school/total,
                  "more than high school but not college": more_than_high_school_not_college/total,
                  "college": college/total}
    
    return dictionary

proportion_of_education()
```

Output:
```
<ipython-input-12-4ac808090666> in proportion_of_education()
      3     # YOUR CODE HERE
      4     # raise NotImplementedError()
----> 5     import pandas as pd
      6     df = pd.read_csv('NISPUF17.csv', index_col=0)
      7     # selecting the needed column from the dataframe

C:\ProgramData\Anaconda3\lib\site-packages\pandas\__init__.py in <module>
    191 # GH 27101
    192 # TODO: remove Panel compat in 1.0
--> 193 if pandas.compat.PY37:
    194 
    195     def __getattr__(name):

AttributeError: partially initialized module 'pandas' has no attribute 'compat' (most likely due to a circular import)
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

Output:
```
<ipython-input-2-2eae51ff42e4> in <module>
----> 1 assert type(proportion_of_education())==type({}), "You must return a dictionary."
      2 assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
      3 assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
      4 assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
      5 assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

<ipython-input-1-80ba6d2b8ec2> in proportion_of_education()
      2     # your code goes here
      3     # YOUR CODE HERE
----> 4     raise NotImplementedError()

NotImplementedError:
```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # YOUR CODE HERE
    # raise NotImplementedError()
    # selecting the needed column from the dataframe
    import pandas as pd
    df = pd.read_csv('NISPUF17.csv', index_col=0)
    # grouping the dataframe by breastfeeding
    df_cpf_01 = df.groupby('CBF_01')
    # gettting the average no. of seasonal influenza for each group
    df_cpf_01_mean = df_cpf_01['P_NUMFLU'].mean()
    
    return (df_cpf_01_mean.iloc[0], df_cpf_01_mean.iloc[1])

average_influenza_doses()
```

Output:
```
<ipython-input-10-8ed3eb650352> in average_influenza_doses()
      3     # raise NotImplementedError()
      4     # selecting the needed column from the dataframe
----> 5     import pandas as pd
      6     df = pd.read_csv('assets/NISPUF17.csv', index_col=0)
      7     # grouping the dataframe by breastfeeding

C:\ProgramData\Anaconda3\lib\site-packages\pandas\__init__.py in <module>
    191 # GH 27101
    192 # TODO: remove Panel compat in 1.0
--> 193 if pandas.compat.PY37:
    194 
    195     def __getattr__(name):

AttributeError: partially initialized module 'pandas' has no attribute 'compat' (most likely due to a circular import)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

Output:
```
<ipython-input-4-380749c0010d> in <module>
----> 1 assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

<ipython-input-3-b22e4dde7c56> in average_influenza_doses()
      1 def average_influenza_doses():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    # raise NotImplementedError()
    import pandas as pd
    df = pd.read_csv('NISPUF17.csv', index_col=0)
    df_sex = df.groupby('SEX')
    # where P_NUMVRC is the total number of varicella doses
    df_male_had_cpox = df['HAD_CPOX'][(df['HAD_CPOX'] == 1) & (df['SEX'] == 1) & (df['P_NUMVRC'] >= 1)].count()
    df_male_not_had_cpox = df['HAD_CPOX'][(df['HAD_CPOX'] == 2) & (df['SEX'] == 1) & (df['P_NUMVRC'] >= 1)].count()
    df_female_had_cpox = df['HAD_CPOX'][(df['HAD_CPOX'] == 1) & (df['SEX'] == 2) & (df['P_NUMVRC'] >= 1)].count()
    df_female_not_had_cpox = df['HAD_CPOX'][(df['HAD_CPOX'] == 2) & (df['SEX'] == 2) & (df['P_NUMVRC'] >= 1)].count()
    return {'male': df_male_had_cpox/df_male_not_had_cpox, 'female': df_female_had_cpox/df_female_not_had_cpox}

chickenpox_by_sex()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

Output:
```
<ipython-input-6-8b72455a9179> in <module>
----> 1 assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

<ipython-input-5-b5b939e2d580> in chickenpox_by_sex()
      1 def chickenpox_by_sex():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    # df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #              "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    # corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    
    import pandas as pd
    df = pd.read_csv('NISPUF17.csv', index_col=0)
    # removing outliers like 77 from (HAD_CPOX) column, and remvong NA values from (P_NUMVRC) column
    df = df[(df['HAD_CPOX'] < 3) & (df['P_NUMVRC'] >= 0)] 
    
    corr, pval=stats.pearsonr(df['HAD_CPOX'], df['P_NUMVRC'])
    
    # just return the correlation
    return corr
    
corr_chickenpox()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

Output:
```
C:\ProgramData\Anaconda3\lib\site-packages\pandas\core\indexes\base.py in <module>
      7 import numpy as np
      8 
----> 9 from pandas._libs import algos as libalgos, index as libindex, lib
     10 import pandas._libs.join as libjoin
     11 from pandas._libs.lib import is_datetime_array

C:\ProgramData\Anaconda3\lib\importlib\_bootstrap.py in parent(self)

KeyboardInterrupt:
```